<a href="https://colab.research.google.com/github/ya22015/CE880_Lab_Task/blob/main/week5/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

---
# Welcome to CE880
### This is your week-5 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
# !pip install -U scikit-learn
import sys
import warnings

if not sys.warnoptions:
    warnings.simplefilter("ignore")
```

---
## Question 1: 
We’re ready to split a larger dataset to solve a regression problem. You’ll use a well-known Boston house prices dataset, which is included in sklearn. This dataset has 506 samples, 13 input variables, and the house values as the output. We have loaded dataset for you and splited into training and test set. Write code to fit linear regression on training `input (x)` and training `output (y)` and print the training score using `model.score` function. 

```python
# from sklearn.datasets import load_boston
from sklearn.datasets import fetch_california_housing
from sklearn.model_selection import train_test_split
import matplotlib.pyplot as plt

california_housing = fetch_california_housing(as_frame=True)
print(california_housing.DESCR)

california_housing.frame.hist(figsize=(12, 10), bins=30, edgecolor="black")
plt.subplots_adjust(hspace=0.7, wspace=0.4)

x = california_housing.data
y = california_housing.target

x_train, x_test, y_train, y_test = train_test_split(x, y, test_size=0.2, random_state=0)
```

<!-- grade_id: my_regression_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression

def my_regression(x_train, y_train):
    """Write code to fit linear regression on training input (x) and training output (y) and 
    print the training score using model.score function. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: my_regression_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_regression(x_train, y_train), 0.60889, rel_tol = 0.05)
```

---
## Question 2: 

Now, implement `RandomForestRegressor` on the same dataset and check the performance on the `test` dataset. Hint: Please read the question carefuly. 

<!-- grade_id: my_RMregression_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestRegressor

def my_RMregression(x_train, y_train, x_test, y_test):
    """Write code to fit RandomForestRegressor on training input (x) and training output (y) and 
    return test score using model.score function. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: my_RMregression_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
t_ans = my_RMregression(x_train, y_train, x_test, y_test)
p_ans = 0.79805
assert math.isclose(t_ans, p_ans, rel_tol = 0.05)
```

---
## Question 3: 
We have loaded wine dataset. Write code to fit LogisticRegression classificaiton model on training data and return the accuracy in percentage % on test data. Please split your dataset into 80% for training and 20% for test

<!-- grade_id: my_classification_ans, points: 0, grade: false, solution: true -->
```python

import numpy as np
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import accuracy_score as accuracy
from sklearn.metrics import classification_report, confusion_matrix
from sklearn.datasets import load_wine
dataset = load_wine()
x = dataset.data
y = dataset.target

def my_classification(x, y):
    """Write code to fit LogisticRegression classificaiton model on training data and 
    return the accuracy in percentage % on test data. 
    Please split your dataset into 80% for training and 20% for test. In logistic regression
    please use solver='liblinear' and random_state=0 """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: my_classification_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
assert math.isclose(my_classification(x, y), 94, rel_tol = 0.05)
```